Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Dynamic Programming - Key Concept
<h4>A problem that can be solved efficiently by DP often has the following properties:</h4>
<ul>
<li>Optimal Substructure(allows recursion): </li>
<ul>
Solution to the problem contains optimal solutions to subproblems.
</ul>
</ul>
<ul>
<li>Overlapping Subproblems(allows speed up): </li>
<ul>
A recursive algorithm revisits the same subproblem over and over again.
</ul>
</ul>


###  <h2>Exercise 3: Hamming Distance(20 points) [4pts/per test (2 hidden tests)]<h2>
<h4>In information theory, the Hamming Distance between two sequences of equal length is the number of positions at which the corresponding symbols are different. </h4>
<h4>Given 2 different sequences, find out the Hamming Distance between 2 sequences. </h4>

<h4> Example: Input sequences "ERROR" and "MINOR" </h4>
<ul>
<li>Output: 3 (<b>ERR</b>OR and <b>MIN</b>OR)</li>
</ul>
<h4> Example: Input sequences "ABBABA" and "ABAAAA" </h4>
<ul>
<li>Output: 2 (AB<b>B</b>A<b>B</b>A and AB<b>A</b>A<b>A</b>A)</li>
</ul>

<!-- grade_id: HA_question, points: 0, grade: false, solution: true -->
```python
def hamming_distance(seq_1, seq_2):
    '''
    hamming_distance():
    Input argument:
        seq_1, seq_2: Tow input sequences (type: str)
    Return: 
        the length of hamming distance between seq_1 into seq_2
    '''
    count = 0
    # YOUR CODE HERE
    raise NotImplementedError()
    return count
```

<!-- grade_id: HA_answer1, points: 4, grade: true, solution: false -->
```python
seq_X = "KATHRIN"
seq_Y = "KERSTIN"
assert hamming_distance(seq_X , seq_Y) == 4
```

<!-- grade_id: HA_answer2, points: 4, grade: true, solution: false -->
```python
seq_X = "1011101"
seq_Y = "1001001"
assert hamming_distance(seq_X , seq_Y) == 2
```

<!-- grade_id: HA_answer3, points: 4, grade: true, solution: false -->
```python
seq_X = "011010100101011"
seq_Y = "010110110101001"
assert hamming_distance(seq_X , seq_Y) == 4
```

<!-- grade_id: HA_answer4, points: 0, grade: true, solution: false -->
```python
#HIDDEN TEST
```

<!-- grade_id: HA_answer5, points: 4, grade: true, solution: false -->
```python
#HIDDEN TEST
```

###  <h2>Exercise 4: Matrix Chain Multiplication(25 points) [5pts/per test (1 hidden test 10pts)]<h2>

<h4> Example: Input matrices</h4>
<ul>
<li>Input: Matrix A row = 10, column = 20</li>
<li>Input: Matrix B row = 20, column = 30</li>
<li>Input: Matrix C row = 30, column = 40</li>
<li>Input: Matrix D row = 40, column = 30</li>
</ul>

<h4>There are many options to multiply a chain of matrices. Ex: (ABC)D = (AB)(CD) = A(BCD) = A(BC)D = ....</h4>
<h4>Cost of ABCD = 10 X 20 X 30 + 10 X 30 X 40 + 10 X 40 X 30 = 30000</h4>
<h4>Cost of A(BC)D = 20 X 30 X 40 + 10 X 20 X 40 + 10 X 40 X 30 = 44000</h4>



<h4> You need to find the most efficient way to multiply these matrices together</h4>
<ul>
<li>Output: 30000 (Method: ((AB)C)D )</li>
</ul>

<!-- grade_id: ma_question, points: 0, grade: false, solution: true -->
```python
import sys
def mcm(p, i, j):
    '''
    mcm(): Matrix Chain Multiplication
    Input argument:
        p: List that descripts the row and column of all matrices(ex: p = [10, 20, 30, 40, 30] implied 4 input matrices with dimension 10X20, 20X30, 30X40, 40X30) (type: list)
        i: Index of p where the subsequence start (type: int)
        j: Index of p where the subsequence end (type: int)
    
    Return: 
        Minimum cost (type: int)
    '''
    minimal_execution = sys.maxsize
    # YOUR CODE HERE
    raise NotImplementedError()
    return minimal_execution
```

<!-- grade_id: MA_ans1, points: 5, grade: true, solution: false -->
```python
matrices_1 = [10, 20, 30, 40, 30]
assert mcm(matrices_1, 1, len(matrices_1)-1) == 30000
```

<!-- grade_id: MA_ans2, points: 5, grade: true, solution: false -->
```python
matrices_2 = [40, 20, 10, 20, 20, 40]
assert mcm(matrices_2, 1, len(matrices_2)-1) == 36000
```

<!-- grade_id: MA_ans3, points: 5, grade: true, solution: false -->
```python
matrices_3 = [30, 10, 10, 40, 30, 60, 10, 50, 10]
assert mcm(matrices_3, 1, len(matrices_3)-1) == 43000
```

<!-- grade_id: MA_ans4, points: 10, grade: true, solution: false -->
```python
#HIDDEN TEST
```

###  <h2>Exercise 5: 0-1 Knapsack Problem(25 points) [5pts/per test (1 hidden test 10pts)] <h2>

<h4> Example: There are a variety of items with different weight and value, ad you have a knapsack of capacity W. Try to get the maximum total value in the knapsack</h4>
<ul>
<li>Input: Item A weight = 10, value = 60</li>
<li>Input: Item B weight = 20, value = 100</li>
<li>Input: Item C weight = 30, value = 120</li>
<li>Input: knapsack W with maximum weight = 50</li>
<li>Output: 220 (Method: B + C)</li>
</ul>

<!-- grade_id: KP_question, points: 0, grade: false, solution: true -->
```python
def knapsack_problem(item_weight, item_value, bag_capacity, order):
  
    '''
    knapsack_problem(): 0-1 Knapsack Problem
    Input argument:
        item_weight: Item's weight (type: list)
        item_value: Item's value (type: list)
        bag_capacity: Capacity of knapsack (type: int)
        order: The item's index that is essencial to the  knapsack # Hint
    Return: 
        Max value (type: int)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: KP_ans1, points: 5, grade: true, solution: false -->
```python
item_value = [40, 50, 60, 70, 80]
item_weight = [5, 10, 15, 20, 25]
bag_capacity = 50
assert knapsack_problem(item_weight, item_value, bag_capacity, len(item_weight)) == 220
```

<!-- grade_id: KP_ans2, points: 5, grade: true, solution: false -->
```python
item_value = [70, 120, 150, 200, 210]
item_weight = [10, 20, 30, 50, 70]
bag_capacity = 100
assert knapsack_problem(item_weight, item_value, bag_capacity, len(item_weight)) == 470
```

<!-- grade_id: KP_ans3, points: 5, grade: true, solution: false -->
```python
item_value = [23, 26, 35, 41, 22]
item_weight = [24, 28, 30, 36, 20]
bag_capacity = 100
assert knapsack_problem(item_weight, item_value, bag_capacity, len(item_weight)) == 102
```

<!-- grade_id: KP_ans4, points: 10, grade: true, solution: false -->
```python
#HIDDEN TEST
```